// Numbas version: finer_feedback_settings
{"name": "Potenser", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": true, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Positive eksponenter", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", ""], "variable_overrides": [[], [], []], "questions": [{"name": "Enkle potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(-9..-2)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{a}^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk at $\\var{a}^2 = \\var{a}\\cdot \\var{a}$</p>"}], "minValue": "a^2", "maxValue": "a^2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(\\var{b})^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Fortegnet er inn i parentesen, som betyr at det er en del av grunntallet:</p>\n<p>$(\\var{b})^2 =(\\var{b})\\cdot (\\var{b})$</p>\n<p>Bruk fortegnsregler.</p>"}], "minValue": "(b)^2", "maxValue": "(b)^2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $-\\var{c}^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Her er fortegnet ikke en del av grunntallet:&nbsp;</p>\n<p>$-{\\var{c}}^2 = -\\var{c}\\cdot \\var{c}$</p>\n<p>Bruk fortegnsregler.</p>\n<p></p>"}], "minValue": "-(c^2)", "maxValue": "-(c^2)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Enkle potenser niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn de f&oslash;lgende potensene.</p>\n<p>$2^0=$[[0]]</p>\n<p>$(-3)^0 =$[[1]]</p>\n<p>$-3^0=$[[2]]</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^0=1$</p>\n<p>Pass p&aring; om fortegnet er en del av grunntallet eller ikke.</p>"}], "gaps": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-2", "-1", "0", "1", "2"], "matrix": [0, 0, 0, "1", 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-3", "-1", "0", "1", "3"], "matrix": [0, 0, 0, "1", 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-3", "-1", "0", "1", "3"], "matrix": [0, "1", 0, "0", 0], "distractors": ["", "", "", "", ""]}], "sortAnswers": false, "inlineCorrectAnswer": true}, {"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn de f&oslash;lgende potensene.</p>\n<p>$2^1 =$[[0]]</p>\n<p>$(-3)^1=$[[1]]</p>\n<p>$-3^1=$[[2]]</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^1=a.$</p>"}], "gaps": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-2", "-1", "0", "1", "2"], "matrix": [0, "0", 0, "0", "1"], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-3", "-1", "0", "1", "3"], "matrix": ["1", 0, 0, "0", 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-3", "-1", "0", "1", "3"], "matrix": ["1", 0, 0, "0", 0], "distractors": ["", "", "", "", ""]}], "sortAnswers": false, "inlineCorrectAnswer": true}, {"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn de f&oslash;lgende potensene.</p>\n<p>$1^6 =$[[0]]</p>\n<p>$(-1)^5=$[[1]]</p>\n<p>$-1^4=$[[2]]</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk at $1^6 = 1 \\cdot 1 \\cdot 1 \\cdot 1 \\cdot 1 \\cdot 1.$</p>\n<p>Pass p&aring; om fortegnet er en del av grunntallet eller ikke.</p>"}], "gaps": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-6", "-1", "0", "1", "6"], "matrix": [0, "0", "0", "1", 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-5", "-1", "0", "1", "5"], "matrix": ["0", "1", 0, "0", 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["-4", "-1", "0", "1", "4"], "matrix": ["0", "1", 0, "0", 0], "distractors": ["", "", "", "", ""]}], "sortAnswers": false, "inlineCorrectAnswer": true}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Produkt av potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om produkt av potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random([2,4,6,8])", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(3..8)", "description": "", "templateType": "anything", "can_override": false}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "random(3..7)", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random(-7..-2)", "description": "", "templateType": "anything", "can_override": false}, "l": {"name": "l", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "o": {"name": "o", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c", "d", "f", "g", "h", "l", "m", "n", "o"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $3^\\var{a} \\cdot 3^{\\var{b}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^n\\cdot a^m= a^{n+m}.$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$3^{\\var{a+b}}$", "$3^{\\var{a b}}$", "$3^{\\var{abs(a-b)}}$", "$3^{\\frac{\\var{a}}{\\var{b}}}$", "$3^{\\var{a+b+3}}$"], "matrix": ["1", 0, 0, 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $2 \\cdot 2^{\\var{c}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk at $2 = 2^1.$</p>\n<p>Husk potensregel: $a^n\\cdot a^m= a^{n+m}.$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$2^{\\var{c+1}}$", "$2^{\\var{c}}$", "$2^{\\var{c+2}}$", "$2^{\\frac{1}{\\var{c}}}$", "$2^{\\var{c-1}}$"], "matrix": ["1", 0, 0, 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(-2)^{\\var{d}}\\cdot (-2)^{\\var{f}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^n\\cdot a^m= a^{n+m}$ -- ogs&aring; n&aring;r $a$ er et negativt tall.</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$(-2)^{\\var{d+f}}$", "$(-2)^{\\var{d * f}}$", "$-2^{\\var{d *f}}$", "$(-2)^{\\var{d-f}}$", "$(-2)^{-\\var{d + f}}$", "$-2^{\\var{d - f}}$", "$-2^{-\\var{d + f}}$"], "matrix": ["1", 0, 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Negative eksponenter", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", ""], "variable_overrides": [[], [], []], "questions": [{"name": "Negative potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(3..9)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Hva er en tilsvarende m&aring;te for &aring; skrive $2^{-\\var{a}}$?</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^{-n} = \\frac{1}{a^n}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$(-2)^{\\var{a}}$", "$2^{\\var{a}}$", "$\\frac{1}{2^{\\var{a}}}$", "$\\frac{1}{(-2)^{\\var{a}}}$", "$\\sqrt[\\var{a}]{2}$", "$\\sqrt[\\var{a}]{-2}$"], "matrix": [0, "0", "1", 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Hva er en tilsvarende m&aring;te for &aring; skrive $\\frac{1}{3^{-\\var{b}}}$?</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $\\frac{1}{a^{-n}} = \\frac{1}{\\frac{1}{a^n}} = 1 : \\frac{1}{a^n} = 1 \\cdot \\frac{a^n}{1} = a^n$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$-3^{\\var{b}}$", "$3^{\\var{b}}$", "$\\frac{1}{3^{\\var{b}}}$", "$\\frac{1}{(-3)^{\\var{b}}}$", "$\\frac{1}{\\sqrt[\\var{b}]{3}}$", "$\\frac{1}{\\sqrt[\\var{b}]{-3}}$"], "matrix": [0, "1", "0", 0, 0, 0], "distractors": ["", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Negative potenser niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>\n<p><strong><em>N&aring;r det gjelder, kan du oppgi svaret som en br&oslash;k.</em></strong></p>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random([-2, -3])", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(-5..-2)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random([-2, -3])", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(-9..9 except -1 except 1 except 0)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c", "d", "f"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{a}^{\\var{b}}.$&nbsp;</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^{-n} = \\frac{1}{a^n}.$</p>"}], "minValue": "a^b", "maxValue": "a^b", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(\\var{c})^{\\var{d}}.$&nbsp;</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^{-n} = \\frac{1}{a^n}.$</p>\n<p>Husk: et negativt tall opph&oslash;yd i</p>\n<ul>\n<li>et oddetall gir et negativt tall;&nbsp;</li>\n<li>et partall gir et positivt tall</li>\n</ul>"}], "minValue": "c^d", "maxValue": "c^d", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{f}^{-1}.$&nbsp;&nbsp;</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $a^{-1} = \\frac{1}{a^1} = \\frac{1}{a}.$</p>"}], "minValue": "1/f", "maxValue": "1/f", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Divisjon av potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(3..7)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(-7..7 except 0 except 1 except -1)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(2..9 except d)", "description": "", "templateType": "anything", "can_override": false}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "random(-7..7 except 0 except 1 except -1)", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random(-7..7 except 0 except 1 except -1 except -g except -(g+c) except g+c)", "description": "", "templateType": "anything", "can_override": false}, "l": {"name": "l", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "o": {"name": "o", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": "100"}, "ungrouped_variables": ["a", "b", "c", "d", "f", "g", "h", "l", "m", "n", "o"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\frac{5^{\\var{d}}}{5^{\\var{f}}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $\\frac{a^n}{a^m} = a^n: a^m = a^{n-m}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$5^{\\var{d-f}}$", "$5^{\\var{f-d}}$", "$5^{\\var{d+f}}$", "$5^{\\var{d f}}$", "$5^{\\frac{\\var{d}}{\\var{f}}}$"], "matrix": ["1", 0, 0, 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\frac{4^{\\var{a}}}{4^{-\\var{b}}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $\\frac{a^n}{a^m} = a^n: a^m = a^{n-m}$.&nbsp;</p>\n<p>Pass p&aring; at n&aring;r eksponenten i nevneren er negativt, f&aring;r du endelig en pluss, for eks:</p>\n<p>$\\frac{2^{15}}{2^{-3}} = 2^{15-(-3)} = 2^{15 + 3} = 2^{18}.$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$4^{\\var{a-b}}$", "$4^{\\var{b-a}}$", "$4^{\\var{a+b}}$", "$4^{\\var{-a*b}}$", "$4^{\\frac{\\var{a}}{-\\var{b}}}$"], "matrix": ["0", 0, "1", 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\frac{3^{\\var{c}}\\cdot 3^{\\var{g}}}{3^{\\var{h}}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Du kan kombinere potensregelen om multiplikasjon og den om divisjon:</p>\n<p>$\\frac{a^n\\cdot a^m}{a^l} = a^{n+m-l}$. Det gjelder b&aring;de med positive og negative ekponenter!</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$3^{\\var{c+g-h}}$", "$3^{\\var{h-g-c}}$", "$3^{\\var{c+g+h}}$", "$3^{\\frac{\\var{c*g}}{\\var{h}}}$", "$3^{-\\var{abs(c*g*h)}}$", "$3^{\\var{abs(c*g*h)}}$", "$3^{\\var{-c-g-h}}$"], "matrix": ["1", "0", 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Potenser og andre regneregler", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", "", ""], "variable_overrides": [[], [], [], []], "questions": [{"name": "Potenser og parentes", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om parentes og potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "random(2..7 except b)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part a", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Part a", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part c", "definition": "clist[indexc]", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Part c", "definition": "random(2..8)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m1": {"name": "m1", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n1": {"name": "n1", "group": "Ungrouped variables", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n2": {"name": "n2", "group": "Part d", "definition": "n2list[indexd]", "description": "", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "Part e", "definition": "random(0..3)", "description": "", "templateType": "anything", "can_override": false}, "alfa1": {"name": "alfa1", "group": "Part e", "definition": "expression(safe(\"(a^2b^3)^3/(a^4b)^2\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa2": {"name": "alfa2", "group": "Part e", "definition": "expression(safe(\"(a^3/b^2)^3a^2b^4\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa3": {"name": "alfa3", "group": "Part e", "definition": "expression(safe(\"a*b^2(a/b^4)^3\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa4": {"name": "alfa4", "group": "Part e", "definition": "expression(safe(\"(a^2b^2)^2(1/(a^3b))^2\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfalist": {"name": "alfalist", "group": "Part e", "definition": "[alfa1, alfa2, alfa3, alfa4]", "description": "", "templateType": "anything", "can_override": false}, "nodd": {"name": "nodd", "group": "Part d", "definition": "random(3,5)", "description": "", "templateType": "anything", "can_override": false}, "njevn": {"name": "njevn", "group": "Part d", "definition": "random(2,4,6)", "description": "", "templateType": "anything", "can_override": false}, "indexd": {"name": "indexd", "group": "Part d", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "n2list": {"name": "n2list", "group": "Part d", "definition": "[njevn, nodd]", "description": "", "templateType": "anything", "can_override": false}, "codd": {"name": "codd", "group": "Part c", "definition": "random(3,5,7,9)", "description": "", "templateType": "anything", "can_override": false}, "cjevn": {"name": "cjevn", "group": "Part c", "definition": "random(2,4,6,8)", "description": "", "templateType": "anything", "can_override": false}, "indexc": {"name": "indexc", "group": "Part c", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "clist": {"name": "clist", "group": "Part c", "definition": "[codd, cjevn]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["f", "m1", "n1"], "variable_groups": [{"name": "Part a", "variables": ["a", "b", "n"]}, {"name": "Part c", "variables": ["c", "d", "codd", "cjevn", "indexc", "clist"]}, {"name": "Part d", "variables": ["n2", "nodd", "njevn", "indexd", "n2list"]}, {"name": "Part e", "variables": ["index", "alfa1", "alfa2", "alfa3", "alfa4", "alfalist"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(\\var{a} \\ \\cdot \\ \\var{b})^{\\var{n}}.$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $(a\\cdot b)^n = a^n\\cdot b^n $</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{a}^{\\var{n}} \\cdot \\var{b}^{\\var{n}}$", "$\\var{a} \\cdot \\var{b}^{\\var{n}}$", "$\\var{a}^{\\var{n}} \\cdot \\var{b}$", "$\\var{a}^{\\var{n+1}} \\cdot \\var{b}^{\\var{n+1}}$", "$\\var{a} \\cdot \\var{b}^{\\var{n+1}}$", "$\\var{a}^{\\var{n+1}} \\cdot \\var{b}$"], "matrix": ["1", 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(-2x)^{\\var{n2}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Du kan tenke p&aring; $-2x$ som produktet av $-2$ og $x$; da:&nbsp;</p>\n<p>$(-2x)^{\\var{n2}} = (-2)^{\\var{n2}}x^{\\var{n2}}$</p>\n<p>Husk: et negativt tall opph&oslash;yd i</p>\n<ul>\n<li>et oddetall gir et negativt tall;&nbsp;</li>\n<li>et partall gir et positivt tall.</li>\n</ul>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{2^{n2}}x^{\\var{n2}}$", "$-\\var{2^{n2}}x^{\\var{n2}}$", "$2x^{\\var{n2}}$", "$-2x^{\\var{n2}}$", "$\\var{2^{n2}}x$", "$-\\var{2^{n2}}x$"], "matrix": ["if(indexd=0, 1,0)", "if(indexd=1, 1,0)", 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(-2)^{\\var {c}}\\cdot 2^{\\var{d}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Du kan tenke p&aring; $(-2)$ som produktet av $-1$ og 2; da&nbsp;</p>\n<p>$(-2)^{\\var{c}}\\cdot 2^{\\var{d}} = (-1\\cdot 2)^{\\var{c}}\\cdot 2^{\\var{d}} =(-1)^{\\var{c}}2^{\\var{c}}\\cdot 2^{\\var{d}} = (-1)^{\\var{c}}2^{\\var{c}+\\var{d}}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$2^{\\var{c+d}}$", "$-2^{\\var{c+d}}$", "$2^{\\var{c}}$", "$-2^{\\var{c}}$", "$2^{\\var{d}}$", "$-2^{\\var{d}}$"], "matrix": ["if(indexc=1,1,0)", "if(indexc=0,1,0)", 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Potenser og br\u00f8k", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser og br&oslash;k.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "random(2..7 except b)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part a", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Part a", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part b", "definition": "random(2..7 except d)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Part b", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Part b", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m1": {"name": "m1", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n1": {"name": "n1", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "Part e", "definition": "random(0..3)", "description": "", "templateType": "anything", "can_override": false}, "alfa1": {"name": "alfa1", "group": "Part e", "definition": "expression(safe(\"(a^2b^3)^3/(a^4b)^2\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa2": {"name": "alfa2", "group": "Part e", "definition": "expression(safe(\"(a^3/b^2)^3a^2b^4\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa3": {"name": "alfa3", "group": "Part e", "definition": "expression(safe(\"a*b^2(a/b^4)^3\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfa4": {"name": "alfa4", "group": "Part e", "definition": "expression(safe(\"(a^2b^2)^2(1/(a^3b))^2\"))", "description": "", "templateType": "mathematical expression", "can_override": false}, "alfalist": {"name": "alfalist", "group": "Part e", "definition": "[alfa1, alfa2, alfa3, alfa4]", "description": "", "templateType": "anything", "can_override": false}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "h*gang", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random(5..9)", "description": "", "templateType": "anything", "can_override": false}, "n2": {"name": "n2", "group": "Ungrouped variables", "definition": "2", "description": "", "templateType": "anything", "can_override": false}, "gang": {"name": "gang", "group": "Ungrouped variables", "definition": "random(4..9)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["g", "h", "n2", "gang"], "variable_groups": [{"name": "Part a", "variables": ["a", "b", "n"]}, {"name": "Part b", "variables": ["c", "d", "m"]}, {"name": "Part c", "variables": ["f", "m1", "n1"]}, {"name": "Part d", "variables": []}, {"name": "Part e", "variables": ["index", "alfa1", "alfa2", "alfa3", "alfa4", "alfalist"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\Big( \\frac{\\var{c}}{\\var{d}} \\Big)^{\\var{m}}.$ </p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $\\left(\\frac{a}{b}\\right) ^n = \\frac{a^n}{b^n}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\frac{\\var{c}^{\\var{m}}}{\\var{d}^{\\var{m}}}$", "$\\frac{\\var{c}}{\\var{d}^{\\var{m}}}$", "$\\frac{\\var{c}^{\\var{m}}}{\\var{d}}$", "$\\frac{\\var{c}^{\\var{m+1}}}{\\var{d}^{\\var{m-1}}}$", "$\\frac{\\var{c}^{\\var{m+1}}}{\\var{d}^{\\var{m+1}}}$", "$\\frac{\\var{c}^{\\var{m-1}}}{\\var{d}^{\\var{m-1}}}$"], "matrix": ["1", 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\Big( \\frac{\\var{a}}{\\var{b}\\cdot \\var{f}} \\Big)^{\\var{n1}}.$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Du kan kombinere potensregler om potenser av br&oslash;k og potenser av produkt:&nbsp;</p>\n<p>$\\left( \\frac{a}{b\\cdot c} \\right)^n = \\frac{a^n}{b^n\\cdot c^n}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\frac{\\var{a}^{\\var{n1}}}{\\var{b}^{\\var{n1}}\\cdot \\var{f}^{\\var{n1}}}$", "$\\frac{\\var{a}}{\\var{b}^{\\var{n1}}\\cdot \\var{f}^{\\var{n1}}}$", "$\\frac{\\var{a}^{\\var{n1}}}{\\var{b}\\cdot \\var{f}}$", "$\\frac{\\var{a}^{\\var{n1}}\\cdot \\var{f}^{\\var{n1}}}{\\var{b}^{\\var{n1}}}$", "$\\frac{\\var{a}^{\\var{n1}}\\cdot \\var{f}^{\\var{n1}}}{\\var{b}}$", "$\\frac{\\var{a}\\cdot \\var{f}}{\\var{b}^{\\var{n1}}}$"], "matrix": ["1", 0, 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\Big(\\frac{\\var{g}}{\\var{h}}\\Big)^{\\var{n2}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Noen ganger er det lurt &aring; forkorte br&oslash;k f&oslash;r vi beregner potensen. Hva er&nbsp; $\\var{g}: \\var{h}$?</p>"}], "minValue": "(g/h)^n2", "maxValue": "(g/h)^n2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Potenser av potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser av potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "random(2..7 except b)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part a", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Part a", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part b", "definition": "random(2..7 except d)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Part b", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Part b", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m1": {"name": "m1", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n1": {"name": "n1", "group": "Part c", "definition": "random(2..7 except m1)", "description": "", "templateType": "anything", "can_override": false}, "m2": {"name": "m2", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "parity": {"name": "parity", "group": "Part a", "definition": "if(mod(m2*n2, 2)=0, 1,0)", "description": "", "templateType": "anything", "can_override": false}, "n2": {"name": "n2", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "m3": {"name": "m3", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "n3": {"name": "n3", "group": "Ungrouped variables", "definition": "random(-9..-2 except -m3)", "description": "", "templateType": "anything", "can_override": false}, "m4": {"name": "m4", "group": "Ungrouped variables", "definition": "random(2,4,6,8)", "description": "", "templateType": "anything", "can_override": false}, "n4": {"name": "n4", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "parity2": {"name": "parity2", "group": "Ungrouped variables", "definition": "if(mod(n4, 2)=0, 1,0)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["m2", "n2", "m3", "n3", "m4", "n4", "parity2"], "variable_groups": [{"name": "Part a", "variables": ["a", "b", "n", "parity"]}, {"name": "Part b", "variables": ["c", "d", "m"]}, {"name": "Part c", "variables": ["f", "m1", "n1"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\big(\\var{f}^{\\var{m1}}\\big)^{\\var{n1}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk potensregel: $(a^m)^n = a^{m\\cdot n}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{f}^{\\var{m1 n1}}$", "$\\var{f}^{\\var{m1 + n1}}$", "$\\var{f}^{\\var{m1 - n1}}$", "$\\var{f}^{\\var{n1-m1}}$", "$\\var{f}^{\\frac{\\var{m1}}{\\var{n1}}}$"], "matrix": ["1", 0, 0, 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\big((-\\var{a})^{\\var{m2}}\\big)^{\\var{n2}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: et negativt tall opph&oslash;yd i</p>\n<ul>\n<li>et oddetall gir et negativt tall;&nbsp;</li>\n<li>et partall gir et positivt tall</li>\n</ul>\n<p>Er $\\var{m2}\\cdot \\var{n2}$ oddetall eller partall?</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{a}^{\\var{m2 n2}}$", "$-\\var{a}^{\\var{m2 + n2}}$", "$\\var{a}^{\\var{m2 + n2}}$", "$-\\var{a}^{\\var{abs(m2 - n2)}}$", "$-\\var{a}^{\\var{m2 n2}}$", "$\\var{a}^{\\var{abs(m2 - n2)}}$"], "matrix": ["if(parity=1,1,0)", 0, 0, 0, "if(parity=0,1,0)", 0], "distractors": ["", "", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Potenser av potenser niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}, {"name": "Oline S\u00e6verud Aksnes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19994/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser av potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Part a", "definition": "random(2..7 except b)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Part a", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Part a", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Part b", "definition": "random(2..7 except d)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Part b", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Part b", "definition": "random(5..8)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m1": {"name": "m1", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "n1": {"name": "n1", "group": "Part c", "definition": "random(2..7)", "description": "", "templateType": "anything", "can_override": false}, "m2": {"name": "m2", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "parity": {"name": "parity", "group": "Part a", "definition": "if(mod(m2*n2, 2)=0, 1,0)", "description": "", "templateType": "anything", "can_override": false}, "n2": {"name": "n2", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "m3": {"name": "m3", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "n3": {"name": "n3", "group": "Ungrouped variables", "definition": "random(-9..-2 except -m3)", "description": "", "templateType": "anything", "can_override": false}, "m4": {"name": "m4", "group": "Ungrouped variables", "definition": "random(2,4,6,8)", "description": "", "templateType": "anything", "can_override": false}, "n4": {"name": "n4", "group": "Ungrouped variables", "definition": "random(3,5,7,9)", "description": "", "templateType": "anything", "can_override": false}, "parity2": {"name": "parity2", "group": "Ungrouped variables", "definition": "if(mod(n4, 2)=0, 1,0)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["m2", "n2", "m3", "n3", "m4", "n4", "parity2"], "variable_groups": [{"name": "Part a", "variables": ["a", "b", "n", "parity"]}, {"name": "Part b", "variables": ["c", "d", "m"]}, {"name": "Part c", "variables": ["f", "m1", "n1"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\big(-\\var{c}^{\\var{m4}}\\big)^{\\var{n4}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Pass p&aring; at n&aring; er fortegnet ikke opph&oslash;ydet i $\\var{m4}.$ Du kan omskrive uttrykket s&aring;nn:</p>\n<p>$\\big(-\\var{c}^{\\var{m4}}\\big)^{\\var{n4}} = \\big(-1 \\cdot \\var{c}^{\\var{m4}}\\big)^{\\var{n4}}= (-1)^{\\var{n4}} \\cdot \\big(\\var{c}^{\\var{m4}}\\big)^{\\var{n4}}$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{c}^{\\var{m4 n4}}$", "$-\\var{c}^{\\var{m4 + n4}}$", "$\\var{c}^{\\var{m4 + n4}}$", "$-\\var{c}^{\\var{abs(m4 - n4)}}$", "$-\\var{c}^{\\var{m4 n4}}$", "$\\var{c}^{\\var{abs(m4 - n4)}}$"], "matrix": ["if(parity2=1,1,0)", 0, 0, 0, "if(parity2=0,1,0)", 0], "distractors": ["", "", "", "", "", ""]}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\big(\\var{b}^{\\var{m3}}\\big)^{\\var{n3}}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk at regneregelen $(a^m)^n = a^{m\\cdot n}$ gjelder ogs&aring; n&aring;r $m$, $n$ eller begge to er negative!&nbsp;</p>\n<p>Husk &aring; bruke fortegnsregler til &aring; bestemme fortegnet til eksponenten $m\\cdot n.$</p>"}], "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "interpretedAnswerForm": "list of list of boolean", "choices": ["$\\var{b}^{\\var{m3* n3}}$", "$-\\var{b}^{\\var{-m3 *n3}}$", "$\\var{b}^{\\var{m3 + n3}}$", "$\\var{b}^{\\var{m3 - n3}}$", "$\\var{b}^{\\frac{\\var{m3}}{\\var{-n3}}}$"], "matrix": ["1", 0, 0, 0, 0], "distractors": ["", "", "", "", ""]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Oppsummering", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": [""], "variable_overrides": [[]], "questions": [{"name": "Oppsummering potenser", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om potenser.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(-9..-2)", "description": "", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(-5..5 except 0 except 1 except -1 except a except b)", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(2..9)", "description": "", "templateType": "anything", "can_override": false}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "random(-5..5 except 0 except 1 except -1 except f except d)", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random(-7..7)", "description": "", "templateType": "anything", "can_override": false}, "l": {"name": "l", "group": "Ungrouped variables", "definition": "random(-7..7 except h)", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "random(-7..7 except h except l)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Ungrouped variables", "definition": "random(-7..7 except 0 except 1 except -1)", "description": "", "templateType": "anything", "can_override": false}, "o": {"name": "o", "group": "Ungrouped variables", "definition": "random(-7..7 except h except l except m)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c", "d", "f", "g", "h", "l", "m", "n", "o"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Utrykk $(3^{\\var{a}}\\cdot 3^{\\var{b}})^{\\var{c}}$ kan omskrives til $3^{x},$ hvor $x$ er et heltall. Skriv $x$ i boksen:</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Denne oppgaven kan l&oslash;ses p&aring; flere m&aring;ter, f.eks: du kan f&oslash;rst summere eksponentene til potensene i parentesen, og s&aring; gange resultatet med den ytre eksponenten (${\\var{c}}$).</p>\n<p>Ellers, kan du gange hver eksponent i parentesen med den ytre eksponenten (${\\var{c}}$), og s&aring; summere de to tallene som du har f&aring;tt.</p>\n<p>Potensregler som du kan bruke (du trenger ikke &aring; bruke alle!) er de f&oslash;lgende:</p>\n<p>$a^n\\cdot a^m = a^{n+m}$</p>\n<p>$(a\\cdot b)^n = a^n \\cdot b^n$</p>\n<p>$(a^m)^n = a^{n+m}$</p>\n<p></p>"}], "minValue": "(a+b)*c", "maxValue": "(a+b)*c", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Utrykk $\\Big(\\frac{2^\\var{d}}{2^{\\var{f}}}\\Big)^{\\var{g}}$ kan omskrives til $2^{x},$ hvor $x$ er et heltall. Skriv $x$ i boksen:</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Denne oppgaven kan l&oslash;ses p&aring; flere m&aring;ter, f.eks: du kan f&oslash;rst bruke potensregelen for br&oslash;k til &aring; beregne eksponenten til 2 inn i parentesen, og s&aring; gange resultatet med den ytre eksponenten (${\\var{g}}$).</p>\n<p>Ellers, kan du gange hver eksponent i parentesen med den ytre eksponenten (${\\var{g}}$), og s&aring; potensregelen for br&oslash;k til &aring; beregne den endelige eksponenten.</p>\n<p>Potensregler som du kan bruke (du trenger ikke &aring; bruke alle!) er de f&oslash;lgende:</p>\n<p>$\\frac{a^n}{a^m} = a^{n-m}$</p>\n<p>$(\\frac{a}{b})^n = \\frac{a^n}{b^n}$</p>\n<p>$(a^m)^n = a^{n+m}$</p>"}], "minValue": "(d-f)*g", "maxValue": "(d-f)*g", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Utrykk $\\Big(\\frac{2^\\var{h}2^{\\var{l}}}{2^{\\var{m}}2^{\\var{o}}}\\Big)^{\\var{n}}$ kan omskrives til $2^{x},$ hvor $x$ er et heltall. Skriv $x$ i boksen:</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Denne oppgaven kan l&oslash;ses p&aring; flere m&aring;ter, f.eks: du kan f&oslash;rst beregne eksponenten inn i parentesen og s&aring; gange resultatet med den ytre eksponenten.</p>\n<p>Ellers, kan du gange hver eksponent i parentesen med den ytre eksponenten, og s&aring; bruk potensreglene for produkt og br&oslash;k til &aring; beregne den endelige eksponenten.</p>\n<p>Potensregler som du kan bruke (du trenger ikke &aring; bruke alle!) er de f&oslash;lgende:</p>\n<p>$a^n\\cdot a^m = a^{n+m}$</p>\n<p>$(a\\cdot b)^n = a^n \\cdot b^n$</p>\n<p>$\\frac{a^n}{a^m} = a^{n-m}$</p>\n<p>$(\\frac{a}{b})^n = \\frac{a^n}{b^n}$</p>\n<p>$(a^m)^n = a^{n+m}$</p>"}], "minValue": "(h+l-m-o)*n", "maxValue": "(h+l-m-o)*n", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "typeendtoleave": false, "startpassword": "", "autoSubmit": true, "allowAttemptDownload": false, "downloadEncryptionKey": "", "showresultspage": "oncompletion"}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"enterreviewmodeimmediately": true, "showactualmarkwhen": "always", "showtotalmarkwhen": "always", "showanswerstatewhen": "always", "showpartfeedbackmessageswhen": "always", "showexpectedanswerswhen": "inreview", "showadvicewhen": "inreview", "allowrevealanswer": true, "intro": "", "end_message": "", "results_options": {"printquestions": true, "printadvice": true}, "feedbackmessages": [], "reviewshowexpectedanswer": true, "showanswerstate": true, "reviewshowfeedback": true, "showactualmark": true, "showtotalmark": true, "reviewshowscore": true, "reviewshowadvice": true}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "extensions": [], "custom_part_types": [], "resources": []}